<a href="https://colab.research.google.com/github/atul-tiwari/Summer_Workshop/blob/master/Atul_Tiwari_LPU_quiz_1.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

## 1(point 1)
Return the size of list.<br>
Example- def list_size([1,2,3,4])-->4

<!-- grade_id: cell-b67bb4ec5c81ce7a, points: 0, grade: false, solution: true -->
```python
def list_size(l):
    '''
    input:
    l: list
    output:
    y: size of the list
    '''
    # YOUR CODE HERE
    y = len(l)
    return y
```

<!-- grade_id: cell-7372255511ce8f91, points: 1, grade: true, solution: false -->
```python
"""Testing code for last element"""
assert list_size([1,2,3,4,5 ]) == 5

```

## 2(point 1)
Return list of last five element of the given list.<br>
Example- def last_five(['a','b','c','d','e','f','g','h'])-->['d','e','f','g','h']

<!-- grade_id: cell-4a9fc6df4c9155de, points: 0, grade: false, solution: true -->
```python
def last_five(l):
    '''
    input:
        l: list
    output:
        y: list with last five elements
    '''
    # YOUR CODE HERE
    y = l[-5:]
    return y

```

<!-- grade_id: cell-2616a801fd497a31, points: 1, grade: true, solution: false -->
```python
"""Testing code for last five elements"""
assert last_five([1,2,3,4,5,6,7]) == [3,4,5,6,7]

```

## 3(point 1)
Check whether given element exists in the list or not. If it exists, then return True otherwise return False.<br>
Example- def check_elem(['a','b','c','d','e','f','g','h'],'a')-->True<br>
         def check_elem(['a','b','c','d','e','f','g','h'],'k')-->False

<!-- grade_id: cell-219fec5b6f23cc54, points: 0, grade: false, solution: true -->
```python
def check_elem(l,elem):
    '''
    Input:
        l: list 
    Output:
        y: boolean in {True, False}
    '''
    # YOUR CODE HERE
    y = False
    if elem in l:
      y =True
    return y
```

<!-- grade_id: cell-6f0ccd65631864c6, points: 1, grade: true, solution: false -->
```python
"""Testing code to check """
assert check_elem([1,2,3,4,5,6,7],10) == False

```

## 4(point 1)
Return the value after multiplication of all the elements of a list.<br>
Example-  mult_value([1,2,3,4,5])-->120

<!-- grade_id: cell-c58c2a108aef10bb, points: 0, grade: false, solution: true -->
```python
def mult_value(l):
    '''
    input:
        l: list
    output:
        y: Return the value after multiplication of all the elements of the list
    '''
    # YOUR CODE HERE
    y = 1 
    for x in l:
      y*=x
    return y

```

<!-- grade_id: cell-c4ca65f6a05c2df3, points: 1, grade: true, solution: false -->
```python
"""Testing code to return the value after Multiplication """
assert mult_value([1,2,3,4]) == 24

```

## 5(point 1)
Return the number of times the element occurs in the list l.<br>
Example- def Element_rept([1,2,3,4,1,2,3,1,1,5],1)-->4<br>
def Element_rept(['a','f','g','h','c','s','d','f','a','b','c','d'],c)-->2

<!-- grade_id: cell-1083c8e7a7fb1b22, points: 0, grade: false, solution: true -->
```python
def Element_rept(l,element):
    '''
    input:
        l: list
        element: element to be checked for
    output:
        count: number of times the element occured in the list
        
    '''
    # YOUR CODE HERE
    count = 0 
    for x in l:
      if x == element:
        count+=1
    return count
    
```

<!-- grade_id: cell-a59073e635b8a564, points: 1, grade: true, solution: false -->
```python
"""Testing code to return the value after Multiplication """
assert Element_rept([1,4,1,2,3,1,5],10) == 0
assert Element_rept(['k','f','g','h','k','s','d','f','a','b','c','d'],'k') == 2
```

## 6(point 1)
Return the size of dictionary.<br>
Example- def dict_size({1:"akshay",2:"vikash"})-->2

<!-- grade_id: cell-fcff3a57231a7c68, points: 0, grade: false, solution: true -->
```python
def dict_size(d):
    '''
    input:
        d: dictionary
    output:
        y: size of the dictionary
    '''
    # YOUR CODE HERE
    y = len(d)
    return y
```

<!-- grade_id: cell-60e0030f95d46962, points: 1, grade: true, solution: false -->
```python
"""Testing code to return the length"""

assert dict_size({1:1,2:2,3:3 }) == 3

```

## 7(point 1)
Return the sorted list of keys.<br>
Examle- def Sort_keyList({5:'IIT K',10:"IIT Jodhpur",3:"IIT Jaipur"})-->[3,5,10]

<!-- grade_id: cell-995343a8f83fc838, points: 0, grade: false, solution: true -->
```python
def Sort_keyList(d):
    '''
    input:
        d: dictionary
    output:
        result: Return the sorted list of keys
        '''
    # YOUR CODE HERE
    result = list(d.keys())
    result.sort()
    return result
```

<!-- grade_id: cell-4c14ad68ef293635, points: 1, grade: true, solution: false -->
```python
"""Testing code to sort the dictionary"""
assert Sort_keyList({5:'IIT K',10:"IIT Jodhpur",3:"IIT Jaipur"}) == [3,5,10]

```

## 8(point 1)
Check wether given number is prime or not. If yes, then return 1, otherwise return 0.<br>
Example- def is_prime(5)-->1<br>
         def is_prime(4)-->0

<!-- grade_id: cell-5b8fe981edaf53ea, points: 0, grade: false, solution: true -->
```python
def is_prime(n):
    '''
    input:
        n: integer
    output:
        return 0 for not prime
        return 1 for prime
        '''
    # YOUR CODE HERE
    if n<=1 :
      return 0
    else:
      for x in range(2,n):
        if n%x==0:
          return 0
      return 1   
```

<!-- grade_id: cell-b7ccc13244089041, points: 1, grade: true, solution: false -->
```python
"""Testing code to check given number is prime or not """
assert is_prime(13) == 1

```

## 9(point 2)
Using the above fuction return list of prime number from a to b.<br>
Example- def primeNumberList(1,10)-->[2, 3, 5, 7]<br>
        

<!-- grade_id: cell-43bbf195e404710e, points: 0, grade: false, solution: true -->
```python
def primeNumberList(a,b):
    '''
    Input:
    a: staring integer
    b: ending integer
    Output:
    return the list of prime numbers from a to b
    '''
    # YOUR CODE HERE
    newl=[]
    for x in range(a,b+1):
      if is_prime(x)==1: # is_prime is function from last question
        newl.append(x)
    return newl
```

<!-- grade_id: cell-7af742941e9ee3ef, points: 1, grade: true, solution: false -->
```python
"""Testing code to list of prime number """
assert primeNumberList(2,10) == [2, 3, 5, 7]
```

```python

```